---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

**---R'da Aritmetik İşlemleri---**
Toplama: +
Çıkarma: -
Çarpma: *
Bölme: /
Üs: ^
Kalan: %%

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
# Toplama: 5 ile 3'ü topla
5+3

# Çıkarma: 8'den 6'yı çıkar

# Çarpma: 3 ile 2'yi çarp

# Bölme: 8'i 2'ye böl

# Üssel: 2 üzeri 3

# Kalan: 9'un 2 ile bölümünden kalan
```

`@solution`
```{r}

```

`@sct`
```{r}

```
